Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
# Please add your name and University of Michagan uniqname here...

NAME = 'Aseem Sachdeva'
UMICH_UNIQNAME = 'aseemsa'
```

---

```python
version = "v2.3.060120"
```

# SIADS 515 Week 1 Homework (HW1)


## A note about the `%%bash` magic command

The lecture material covered integrating Linux shell commands with Jupyter notebooks using the !, !!, %sx, or %system magic commands. There's one more integration to introduce to you: the %%bash magic command. bash is a popular linux shell, and starting a Jupyter notebook cell with the %%bash magic command will tell Jupyter that the remainder of the cell is to be interpreted as linux (bash) commands.


## A note about the `tee` shell command

For the shell questions in this assignment, you'll generally need to redirect your output to a specifc file for the autograder cells to work. The `tee` command will come in handy, it takes input from STDIN and writes it to STDOUT *and* a file that you specify. This is helpful because you'll be able to see your answer and the file will be writen to allow the autograder to do its work.

## Example Question

Copy the first 3 lines of `assets/ride.csv` into a file named `solution_example.txt`

*(Note that this is an example problem to demonstrate how to answer shell-based questions in this assignment.)*

```python
%%bash

head -3 assets/ride.csv | tee solution_example.txt
```

Output:
```
﻿Type,Local Number,Message,Field 1,Value 1,Units 1,Field 2,Value 2,Units 2,Field 3,Value 3,Units 3,Field 4,Value 4,Units 4,Field 5,Value 5,Units 5,Field 6,Value 6,Units 6,Field 7,Value 7,Units 7,Field 8,Value 8,Units 8,Field 9,Value 9,Units 9,Field 10,Value 10,Units 10,Field 11,Value 11,Units 11,Field 12,Value 12,Units 12,Field 13,Value 13,Units 13,Field 14,Value 14,Units 14,Field 15,Value 15,Units 15,Field 16,Value 16,Units 16,Field 17,Value 17,Units 17,Field 18,Value 18,Units 18,Field 19,Value 19,Units 19,Field 20,Value 20,Units 20,Field 21,Value 21,Units 21,Field 22,Value 22,Units 22,Field 23,Value 23,Units 23,Field 24,Value 24,Units 24,Field 25,Value 25,Units 25,Field 26,Value 26,Units 26,Field 27,Value 27,Units 27,Field 28,Value 28,Units 28,Field 29,Value 29,Units 29,Field 30,Value 30,Units 30,Field 31,Value 31,Units 31,Field 32,Value 32,Units 32,Field 33,Value 33,Units 33,Field 34,Value 34,Units 34,Field 35,Value 35,Units 35,Field 36,Value 36,Units 36,Field 37,Value 37,Units 37,F
```

<!-- grade_id: cell-1d6b28764fdbf76b, points: 0, grade: true, solution: false -->
```python
# Example Question, Check 1
# 0 points

with open('solution_example.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 3, 'solution file must be 3 lines long'

# There are no hidden autograder tests in this cell.
```

## Question 1
Copy the last 5 lines of `assets/ride.csv` into a file called `solution01.txt`.

<!-- grade_id: cell-d2f755630d46555d, points: 0, grade: false, solution: true -->
```python
%%bash

tail -5 assets/ride.csv | tee solution01.txt

# YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
Data,10,lap,timestamp,"896019812",s,start_time,"896019530",,start_position_lat,"504490526",semicircles,start_position_long,"-998989384",semicircles,end_position_lat,"504426837",semicircles,end_position_long,"-999050457",semicircles,total_elapsed_time,"276.443",s,total_timer_time,"276.443",s,total_distance,"981.77",m,unknown,"504490526",,unknown,"-998989384",,unknown,"504424374",,unknown,"-999051451",,message_index,"1",,total_calories,"55",kcal,avg_speed,"3.551",m/s,max_speed,"5.496",m/s,total_ascent,"24",m,total_descent,"1",m,event,"9",,event_type,"1",,lap_trigger,"7",,sport,"2",,sub_sport,"0",,enhanced_avg_speed,"3.551",m/s,enhanced_max_speed,"5.496",m/s,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,
Definition,12,session,timestamp,1,,start_time,1,,start_position_lat,1,,start_position_long,1,,total_elapsed_time,1,,total_timer_time,1,,total_distance,1,,total_cycles,1,,nec_lat,1,,nec_
```

```python
import csv
import shutil
import fileinput

##Import some libraries and set some variables for question 6

to_add = ['timestamp','latitude','longitude','distance','altitude,speed']

my_file_name = "assets/ride.csv"
my_file_name_2 = "cleaned_file1.csv"
my_file_name_3 = "cleaned_file2.csv"
my_file_name_4 = "cleaned_file3.csv"
my_file_name_5 = "cleaned_file4.csv"
my_file_name_6 = "cleaned_file5.csv"
my_file_name_7 = "cleaned_file6.csv"
my_file_name_8 = "cleaned_file7.csv"
my_file_name_9 = "cleaned_file8.csv"
my_file_name_10 = "cleaned_file9.csv"
my_file_name_11 = "cleaned_file10.csv"
```

```python
cleaned_file = "cleaned_file1.csv"
cleaned_file_2 = "cleaned_file2.csv"
cleaned_file_3 = "cleaned_file3.csv"
cleaned_file_4 = "cleaned_file4.csv"
cleaned_file_5 = "cleaned_file5.csv"
cleaned_file_6 = "cleaned_file6.csv"
cleaned_file_7 = "cleaned_file7.csv"
cleaned_file_8 = "cleaned_file8.csv"
cleaned_file_9 = "cleaned_file9.csv"
cleaned_file_10 = "cleaned_file10.csv"
cleaned_file_11 = "cleaned_file11.csv"

```

<!-- grade_id: cell-1e37b44d1c1726b5, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 1
# 1 point

with open('solution01.txt') as ans:
    lines = ans.readlines()
    
assert len(lines) == 5, 'solution file must be 5 lines long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-dd61bcafaa99e567, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 2
# 1 point

# There are hidden autograder tests in this cell.
```

## Question 2
Copy lines 15 through 20 from `assets/ride.csv` into a file called `solution02.txt`.

<!-- grade_id: cell-018ef04758dd9372, points: 0, grade: false, solution: true -->
```python
%%bash

sed -n '15,20 p' assets/ride.csv > solution02.txt



# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-450c38d1d8739c1b, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 1
# 1 point

with open('solution02.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 6, 'solution file must be 6 lines long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-0e662ecb874254bf, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 2
# 1 points

# There are hidden autograder tests in this cell.
```

## Question 3
Copy all lines from `assets/ride.csv` that contain "Data,10", "Data,11", "Data,12", or "Data,13" into a file called `solution03.txt`. The lines in the solution should be in the same order as they appear in the ride.csv file.

<!-- grade_id: cell-da32244325922c68, points: 0, grade: false, solution: true -->
```python
%%bash


grep -e "Data,10" -e "Data,11" -e "Data,12" -e "Data,13" assets/ride.csv > solution03.txt


# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-5d6d674650024aa9, points: 1, grade: true, solution: false -->
```python
# Question 3, Check 1
# 1 point

with open('solution03.txt') as ans:
    lines = ans.readlines()

assert len(lines) > 1, 'solution file must have more than one line'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-9d55e0dd789cd561, points: 1, grade: true, solution: false -->
```python
# Question 3, Check 2
# 1 point

# There are hidden autograder tests in this cell.
```

### Question 4

Count the lines from `assets/ride.csv` that contain "Data,9". Create a file called `solution04.txt` that has one line that contains the number of lines in ride.csv. That line should not contain anything other than the number of lines from 

<!-- grade_id: cell-0cf2e81e963c5402, points: 0, grade: false, solution: true -->
```python
%%bash


 grep -o "Data,9" assets/ride.csv | wc -l > solution04.txt



# Remove the "raise NotImplemented" placeholder Python code below. This is an artifact 
# of the autograder, even though this is supposed to be a "%%bash" cell.

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-e175ee39f82440ef, points: 1, grade: true, solution: false -->
```python
# Question 4, Check 1
# 1 point

with open('solution04.txt') as ans:
    lines = ans.readlines()

assert len(lines) == 1, 'solution file should be 1 line long'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-036ce5a0a73e018b, points: 1, grade: true, solution: false -->
```python
# Question 4, Check 2
# 1 point

# There are hidden autograder tests in this cell.

# Checking if line 1 of solution04.txt ONLY contains the correct answer.
```

## Three solutions to the same problem...

For Question 5, 6, 7, you will create a CSV file for each with the same content based on the data from `assets/ride.csv`. You'll use a different solution for each problem:
  - Question 5: with shell commands
  - Question 6: with Python (without Pandas)
  - Question 7: with Pandas


You'll notice that although the `ride.csv` file is a csv file, it looks a bit different than most csv files.  The first few columns are always of the form (using regular expression notation) ```"(Data|Definition),[\d]+,record"```.  After those columns, the data appear as a key-value-unit triple.  For example:

```
Data,9,record,timestamp,"896018545",s,position_lat,"504719750",semicircles,position_long,"-998493490",semicircles,distance,"10.87",m,altitude,"285.79999999999995",m,speed,"1.773",m/s,unknown,"3929",,unknown,"1002",,enhanced_altitude,"285.79999999999995",m,enhanced_speed,"1.773",m/s,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,,
```

In this case, you have:

```
timestamp,"896018545",s
position_lat,"504719750",semicircles
position_long,"-998493490",semicircles
speed,"1.773",m/s
```


The data we are interested in comes from the lines in `assets/ride.csv` that start with "Data,9,record". 

For each question, you will create a CSV file with these headers, and this first file:

| timestamp| latitude| longitude| distance|          altitude| speed|
|----------|---------|----------|---------|------------------|------|
| 896018545|504719750|-998493490|    10.87|285.79999999999995| 1.773|
|       ...|      ...|       ...|      ...|               ...|   ...|
|       ...|      ...|       ...|      ...|               ...|   ...|

## Question 5 - Shell

Create a new file that matches the format described above.
  - The file should be named `solution05.csv`
  - Solve the problem using only Linux shell commands

<!-- grade_id: cell-3ddf21c4bdc70e67, points: 0, grade: false, solution: true -->
```python
%%bash
rm -f solution05.csv


grep -e "Data,9" assets/ride.csv  > intermediate_solution1.csv

sed 's/Data,9,record,timestamp,//' intermediate_solution1.csv > intermediate_solution2.csv

sed 's/s,//' intermediate_solution2.csv > intermediate_solution3.csv

sed 's/position_lat,//' intermediate_solution3.csv > intermediate_solution4.csv

sed 's/position_long,//' intermediate_solution4.csv > intermediate_solution5.csv

sed 's/distance,//' intermediate_solution5.csv > intermediate_solution6.csv

sed 's/altitude,//' intermediate_solution6.csv > intermediate_solution7.csv

sed 's/semicircles,//' intermediate_solution7.csv > intermediate_solution8.csv

sed 's/speed,//' intermediate_solution8.csv > intermediate_solution9.csv

cut -d',' -f1-3,5- intermediate_solution9.csv > intermediate_solution10.csv

cut -d',' -f1-4,6- intermediate_solution10.csv > intermediate_solution11.csv

cut -d',' -f1-5,7- intermediate_solution11.csv > intermediate_solution12.csv

cut -d',' -f1-6,8- intermediate_solution12.csv > intermediate_solution13.csv

cut -d',' -f1-6 intermediate_solution13.csv > intermediate_solution14.csv

sed -i 1i'timestamp,latitude,longitude,distance,altitude,speed' intermediate_solution14.csv 


cp intermediate_solution14.csv solution05.csv

```

```python
with open('solution05.csv') as ans:
        ans_reader = csv.reader(ans, delimiter=',', quotechar='"')
        header = next(ans_reader)
        print(header[0])
        print(header[1])
        print(header[2])
        print(header[3])
        print(header[4])
        print(header[5])
```

Output:
```
timestamp
latitude
longitude
distance
altitude
speed
```

<!-- grade_id: cell-7652086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 5, Check 1
# 1 point

import csv

def close_match(a, b):
    return round(float(a), 3) == round(float(b), 3)

def ride_csv_check1(filename):
    with open(filename) as ans:
        ans_reader = csv.reader(ans, delimiter=',', quotechar='"')

        header = next(ans_reader)
        assert header[0] == 'timestamp', 'incorrect header'
        assert header[1] == 'latitude', 'incorrect header'
        assert header[2] == 'longitude', 'incorrect header'
        assert header[3] == 'distance', 'incorrect header'
        assert header[4] == 'altitude', 'incorrect header'
        assert header[5] == 'speed', 'incorrect header'

        first_line = next(ans_reader)
        # checks values to three decimal places
        assert close_match(first_line[0], 896018545), 'incorrect value'
        assert close_match(first_line[1], 504719750), 'incorrect value'
        assert close_match(first_line[2], -998493490), 'incorrect value'
        assert close_match(first_line[3], 10.87), 'incorrect value'
        assert close_match(first_line[4], 285.79999999999995), 'incorrect value'
        assert close_match(first_line[5], 1.773), 'incorrect value'

        
solution_filename = 'solution05.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-4f3a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 5, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-c179efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 5, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```

## Q06 - Python without Pandas

Create a new file that matches the format described above:
  - The file should be named `solution06.csv`
  - Don't use any shell or Jupyter magic commands
  - Solve the problem using Python
  - You may use the Python standard library (e.g. the built-in [csv](https://docs.python.org/3/library/csv.html) package)
  - You may NOT use third-party libraries (e.g. Pandas)

```python
%%bash
rm -f *.csv
```

<!-- grade_id: cell-e53a157e0c081aba, points: 0, grade: false, solution: true -->
```python
with open(my_file_name, 'r', newline='') as infile, \
     open(cleaned_file, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        if line[1] == '9' and line[0] != 'Definition' :
            writer.writerow(line)
      
   


              
with open(my_file_name_2, 'r', newline='') as infile, \
     open(cleaned_file_2, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[0:4]
        writer.writerow(line)
        

with open(my_file_name_3, 'r', newline='') as infile, \
     open(cleaned_file_3, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[1:2]
        writer.writerow(line)
            

            
with open(my_file_name_4, 'r', newline='') as infile, \
     open(cleaned_file_4, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[1]
        writer.writerow(line)
        
with open(my_file_name_5, 'r', newline='') as infile, \
     open(cleaned_file_5, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[2:3]
        writer.writerow(line)         
        
with open(my_file_name_6, 'r', newline='') as infile, \
     open(cleaned_file_6, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[2]
        writer.writerow(line)   
        
with open(my_file_name_7, 'r', newline='') as infile, \
     open(cleaned_file_7, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[3:4]
        writer.writerow(line)       
        
        
with open(my_file_name_8, 'r', newline='') as infile, \
     open(cleaned_file_8, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[3]
        writer.writerow(line)      
        
with open(my_file_name_9, 'r', newline='') as infile, \
     open(cleaned_file_9, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[4:6]
        writer.writerow(line)      

with open(my_file_name_10, 'r', newline='') as infile, \
     open(cleaned_file_10, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[5:7]
        writer.writerow(line)    
        
        
with open(my_file_name_11, 'r', newline='') as infile, \
     open(cleaned_file_11, 'w',newline='') as outfile:
    writer = csv.writer(outfile)
    for line in csv.reader(infile, delimiter=','):
        del line[6:]
        writer.writerow(line)  

for line in fileinput.input(files=['cleaned_file11.csv'], inplace=True):
    if fileinput.isfirstline():
        print('timestamp,latitude,longitude,distance,altitude,speed')
    print(line)
    
with open('cleaned_file11.csv') as input, open('solution06.csv', 'w') as output:
    non_blank = (line for line in input if line.strip())
    output.writelines(non_blank)

#raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-q652086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 6, Check 1
# 1 point

# note that this function is defined in the cell for Question 5, Check 1 above
solution_filename = 'solution06.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-q63a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 6, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-q679efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 6, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```

## Q07 -  Pandas

Create a new file that matches the format described above:
  - The file should be named `solution07.csv`
  - Don't use any shell or Jupyter magic commands
  - Solve the problem using functionality provided by the Pandas library in Python

<!-- grade_id: cell-184c965f0270227f, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
pd.set_option('display.max_columns', None)


ride_copy = pd.read_csv('assets/ride.csv')

ride_copy_2 = ride_copy.iloc[22:]

ride_copy_2.head()

ride_copy_3 = ride_copy_2[ride_copy_2['Local Number'] == 9]

ride_copy_3.drop("Type", axis=1, inplace=True)

ride_copy_3.drop("Local Number", axis=1, inplace=True)

ride_copy_3.drop("Message", axis=1, inplace=True)

ride_copy_3.drop("Field 1", axis=1, inplace=True)

ride_copy_3.drop("Units 1", axis=1, inplace=True)

ride_copy_3.drop("Field 2", axis=1, inplace=True)

ride_copy_3.drop("Units 2", axis=1, inplace=True)

ride_copy_3.drop("Field 3", axis=1, inplace=True)

ride_copy_3.drop("Units 3", axis=1, inplace=True)

ride_copy_3.drop("Field 4", axis=1, inplace=True)

ride_copy_3.drop("Units 4", axis=1, inplace=True)

ride_copy_3.drop("Field 5", axis=1, inplace=True)

ride_copy_3.drop("Units 5", axis=1, inplace=True)

ride_copy_3.drop("Field 6", axis=1, inplace=True)


ride_copy_3.drop(ride_copy_3.columns[6:], axis=1, inplace=True)

ride_copy_3 = ride_copy_3.rename(columns={'Value 1': 'timestamp', 'Value 2': 'latitude', 'Value 3': 'longitude', 'Value 4': 'distance', 'Value 5': 'altitude', 'Value 6': 'speed'})

ride_copy_3.reset_index(inplace=True)

ride_copy_3['latitude'] = ride_copy_3['latitude'].astype(np.int64)

ride_copy_3['longitude'] = ride_copy_3['longitude'].astype(np.int64)

ride_copy_3.drop("index", axis=1, inplace=True)

ride_copy_3.to_csv('solution07.csv', index = False)


#ride_copy_3.head()


# with open(my_file_name_11, 'r', newline='') as infile, \
#      open(cleaned_file_11, 'w',newline='') as outfile:
#     writer = csv.writer(outfile)
#     for line in csv.reader(infile, delimiter=','):
#         del line[6:]
#         writer.writerow(line)  

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-q752086039bcee15, points: 1, grade: true, solution: false -->
```python
# Question 7, Check 1
# 1 point

# note that this function is defined in the cell for Question 5, Check 1 above
solution_filename = 'solution07.csv'
ride_csv_check1(solution_filename)
# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-q73a777caf2655b5, points: 1, grade: true, solution: false -->
```python
# Question 7, Check 2
# 1 point

# There are hidden tests in this cell to check the solution has the correct number of rows.
```

<!-- grade_id: cell-q779efeff3149eea, points: 12, grade: true, solution: false -->
```python
# Question 7, Check 3
# 12 points

# There are hidden tests in this cell to check the solution has all the correct values for each column.
```

```python

```